---
rule_files:
  - "alerts.rule"
evaluation_interval: 1m
tests:
  # MissingCheckpoint
  - interval: 5m
    input_series:
      - series: >-
          polygon_validators_monitoring_missed_checkpoints{moniker="Tank",
          vid="42", checkpoint="2"}
        values: _ 1 1 1 _ _ _
      - series: >-
          polygon_validators_monitoring_mentions{vid="42", mentions="superman"}
        values: 1+0x10
    alert_rule_test:
      # event should fire
      - eval_time: 5m
        alertname: MissingCheckpoint
        exp_alerts:
          - exp_labels:
              severity: high
              moniker: Tank
              vid: 42
              mentions: superman
              checkpoint: 2
            exp_annotations:
              summary: Validator Tank missed checkpoint 2!
      # should stop
      - eval_time: 25m
        alertname: MissingCheckpoint

  # MultiOperatorMissingCheckpoint
  - interval: 5m
    input_series:
      - series: >-
          polygon_validators_monitoring_missed_checkpoints{moniker="Panther",
          vid="42", checkpoint="2"}
        values: _ 1 1 1 _ _ _
      - series: >-
          polygon_validators_monitoring_missed_checkpoints{moniker="Tiger",
          vid="73", checkpoint="2"}
        values: _ 1 1 1 _ _ _
    alert_rule_test:
      # event should fire
      - eval_time: 5m
        alertname: MultiOperatorMissingCheckpoint
        exp_alerts:
          - exp_labels:
              severity: high
              checkpoint: 2
            exp_annotations:
              summary: At least 2 validators missed the checkpoint 2!
      # should stop
      - eval_time: 25m
        alertname: MissingCheckpoint

  # ConsecutiveMissingCheckpoint
  - interval: 5m
    input_series:
      - series: >-
          polygon_validators_monitoring_missed_checkpoints_in_row{moniker="Panther", vid="42"}
        values: _ 5x10
    alert_rule_test:
      # event should fire
      - eval_time: 5m
        alertname: ConsecutiveMissingCheckpoint
        exp_alerts:
          - exp_labels:
              severity: critical
              moniker: Panther
              vid: 42
            exp_annotations:
              summary: Validator Panther missed 5 checkpoints in a row!
      # should stop
      - eval_time: 25m
        alertname: ConsecutiveMissingCheckpoint

  # MultiOperatorConsecutiveMissingCheckpoint
  - interval: 5m
    input_series:
      - series: >-
          polygon_validators_monitoring_missed_checkpoints_in_row{moniker="Panther", vid="42"}
        values: _ 5x10
      - series: >-
          polygon_validators_monitoring_missed_checkpoints_in_row{moniker="Tiger", vid="73"}
        values: _ 3x10
    alert_rule_test:
      # event should fire
      - eval_time: 5m
        alertname: MultiOperatorConsecutiveMissingCheckpoint
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: 2 validators missed more than 3 checkpoints in a row!
      # should stop
      - eval_time: 25m
        alertname: MultiOperatorConsecutiveMissingCheckpoint

  # ValidatorInGP
  - interval: 5m
    input_series:
      - series: >-
          polygon_validators_monitoring_under_pb_strike{moniker="Tank"}
        values: 0 1 700+700x3
    alert_rule_test:
      # GP1
      - eval_time: 5m
        alertname: ValidatorEntersGP
        exp_alerts:
          - exp_labels:
              severity: high
              moniker: Tank
            exp_annotations:
              summary: Validator Tank enters Grace Period!
      # Mute alerts not firing yet
      - eval_time: 5m
        alertname: ValidatorEntersGPMute
      # After 5 minutes mute should exist
      - eval_time: 10m
        alertname: ValidatorEntersGPMute
        exp_alerts:
          - exp_labels:
              moniker: Tank
      # GP2
      - eval_time: 10m
        alertname: ValidatorGPBump
        exp_alerts:
          - exp_labels:
              severity: high
              moniker: Tank
            exp_annotations:
              summary: Validator Tank is in Grace Period 2!
      # FN only
      - eval_time: 15m
        alertname: ValidatorGPBump
        exp_alerts:
          - exp_labels:
              severity: critical
              moniker: Tank
            exp_annotations:
              summary: >-
                Validator Tank has received Final Notice
                and will be unstaked soon.
      # should keep ValidatorEntersGP till recover
      - eval_time: 20m
        alertname: ValidatorEntersGP
        exp_alerts:
          - exp_labels:
              severity: high
              moniker: Tank
            exp_annotations:
              summary: Validator Tank enters Grace Period!

# vim: set ts=2 sw=2 ft=yaml:
